## DESCRIPTION
## Linear Algebra
## ENDDESCRIPTION

## Tagged by cmd6a 5/3/06

## DBsubject(Linear algebra)
## DBchapter(Eigenvalues and eigenvectors)
## DBsection(Computing eigenvalues and eigenvectors)
## Date(July 2013)
## Institution(Rochester and Hope College)
## Author(Paul Pearson)
## Level(3)
## MO(1)
## TitleText1('Linear Algebra with Applications')
## AuthorText1('Jeffrey Holt')
## EditionText1('1')
## Section1('8.3')
## Problem1('')

DOCUMENT();        # This should be the first executable line in the problem.

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());
$showPartialCorrectAnswers = 1;

Context('Matrix');
Context()->flags->set(
  tolerance => 0.001,
  tolType => "absolute",
);

$case = random(5,10,5);

if ($case==5) {
	$u[1][1] = random(1,2,1);
	$u[2][1] = 3 - $u[1][1];
} else {
        $u[1][1] = random(1,3,2);
        $u[2][1] = 4 - $u[1][1];  
}

$u[1][1] = $u[1][1] * random(-1,1,2);
$u[2][1] = $u[2][1] * random(-1,1,2);
$u[1][2] = - $u[2][1]; 
$u[2][2] = $u[1][1]; 

$U = Matrix([
[ $u[1][1], $u[1][2] ],
[ $u[2][1], $u[2][2] ],
]);

$eig1 = $case * non_zero_random(-5,-1,1) * random(0,1,1);
if ($eig1 == 0) { 
	$eig2 = $case * non_zero_random(1,5,1);
} else { 
	$eig2 = 0;
}

$E = Matrix([
[$eig1, 0],
[0, $eig2],
]);

$I = $U->inverse;

$A = $U * $E * $I;

$lambda[1] = Compute("$eig1");
$lambda[2] = Compute("$eig2");

$basis[1] = Matrix([ [$u[1][1]], [$u[2][1]] ]);
$basis[2] = Matrix([ [$u[1][2]], [$u[2][2]] ]);

Context()->texStrings;
BEGIN_TEXT
Find the eigenvalues of the matrix 
\( A = $A \)
and a corresponding eigenvector for each eigenvalue.
Scale the eigenvector, if necessary, to have integer entries (no fractions).
$HR
The smaller eigenvalue is \{ $lambda[1]->ans_rule(10) \} 
and has eigenvector \( \vec{v}_1 = \) \{ $basis[1]->ans_array(15) \}.
$BR
$BR
The larger eigenvalue is \{ $lambda[2]->ans_rule(10) \} 
and has eigenvector \( \vec{v}_2 = \) \{ $basis[2]->ans_array(15) \}.
$BR
END_TEXT
Context()->normalStrings;

foreach my $i (1..2) {

ANS($lambda[$i]->cmp);
ANS($basis[$i]->cmp(checker=>sub {
  my ( $correct, $student, $ansHash ) = @_;
  return (($student == $correct) or ($student == -1*$correct));
}));

}

COMMENT('One of the eigenvalues is zero and the other is nonzero.');
ENDDOCUMENT();